---
title: "Feedback at scale"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
    css: "css/font-size.css"
runtime: shiny_prerendered
---

<!-- Add JavaScript code for making the exercise code larger -->
<script language="JavaScript" src="js/exercise-font-size.js"></script>

```{r setup, include=FALSE}
# packages
library(learnr)
library(tidyverse)
#devtools::install_github("rstudio-education/gradethis")
library(gradethis)
#devtools::install_github("rundel/learnrhash")
library(learnrhash)
library(glue)

# tutorial options
tutorial_options(
  # code running in exercise times out after 30 seconds
  exercise.timelimit = 30, 
  # use gradethis for checking
  exercise.checker = gradethis::grade_learnr
  )

# hide non-exercise code chunks
knitr::opts_chunk$set(
  echo = FALSE,
  warning = FALSE,
  message = FALSE
  )
```

---

## About me

```{r stormtrooper, fig.margin = TRUE, echo = FALSE, out.width = "100%", fig.cap="Looking at the Edinburgh castle from Arthur's seat."}
knitr::include_graphics("images/edinburgh.png")
```

```{r about-me}
question(
  "Which of the following is true about me? Check all that apply.",
  answer("I'm part of the education team at RStudio.", correct = TRUE),
  answer("I am also a Senior Lecturer of Statistics and Data Science at the University of Edinburgh.", correct = TRUE),
  answer("I have four dogs."), 
  allow_retry = TRUE,
  incorrect = "Not quite. I love dogs, but I don't have any! I have four cats.",
  correct = "Correct! And one more tidbit: I love R!"
  )
```

## Where are we?

```{r tada, exercise=TRUE}
glue(
  '{emo::ji("tada")} - Enrolments in stats & data science courses are increasing \n',
  '{emo::ji("smile")} - These courses are getting more computational \n',
  '{emo::ji("face_with_head_bandage")} - Providing timely and meaningful feedback at scale is hard \n',
  '{emo::ji("cold_sweat")} - Doing so with fully online delivery of courses is even harder',
  )
```

## Where do we go from here?

```{r load-package, exercise = TRUE}
library(learnr)     # interactive exercises
library(gradethis)  # automated feedback
library(dsbox)      # tutorials in a package
library(learnrhash) # collecting student data
```

```{r load-package-solution}
library(learnr)     # interactive exercises
library(gradethis)  # automated feedback
library(dsbox)      # tutorials in a package
library(learnrhash) # collecting student data
```

```{r load-package-check}
grade_code(
  correct = "Packages are loaded, and we're ready to roll!"
)
```
